<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Einführung Machine Learning
### Sommersemester 2026
Prof. Dr. Heiner Giefers

## Aufgabe: Grundlegende Statistik in Python

In dieser Aufgabe verwenden wir die Python-Bibliotheken **NumPy** und **Matplotlib**, um eine Stichprobe zu erzeugen, zu analysieren und und zu visualisieren.

Wir verwenden ein realitätsnahes Beispiel – die **Körpergröße von Personen** – und berechnen wichtige statistische Kennzahlen wie **Mittelwert**, **Median** und die **Standardabweichung**. Anschließend stellen wir die Verteilung visuell dar und interpretieren Ihre Ergebnisse.

```python
import numpy as np
import matplotlib.pyplot as plt
```

**Aufgabe: Stichprobe erzeugen**

Erzeugen Sie eine Zufallsstichprobe mit 500 Werten aus einer **Normalverteilung** mit dem Mittelwert \( \mu = 170 \) und der Standardabweichung \( \sigma = 8 \). Verwenden Sie dafür die Funktion `np.random.normal`.

Weisen Sie die generierten Werte der Variable `data` zu.

<!-- grade_id: cell-8de9dd2a49621df5, points: 0, grade: false, solution: true -->
```python
import numpy as np
data = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-79aaddc1e8771e16, points: 0, grade: true, solution: false -->
```python
# Test Zelle
assert 'data' in globals(), "Variable 'data' wurde nicht definiert"
assert isinstance(data, np.ndarray), "'data' muss ein NumPy-Array sein"
assert len(data) == 500, "Die Stichprobe muss genau 500 Werte enthalten"
```

**Aufgabe: Statistische Kennzahlen**

Berechnen Sie auf Basis der generierten Stichprobe `data` folgende Kennzahlen:

- Arithmetisches Mittel (mean)
- Median
- Varianz
- Standardabweichung
- Interquartilsabstand (IQR)

Speichern Sie die Ergebnisse in den Variablen `mean_val`, `median_val`, `var_val`, `std_val` und `iqr_val`.

```python
mean_val = None    # Variable für den Mittelwert
median_val = None  # Variable für den Median
var_val = None     # Variable für die Varianz
std_val = None     # Variable für die Standardabweichungd

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Test Zelle
assert all(name in globals() for name in [
    "mean_val", "median_val", "var_val", "std_val"
]), "Mindestens eine der geforderten Variablen fehlt."
assert isinstance(mean_val, (float, np.float64)), "Mittelwert sollte ein float sein."

```

**Aufgabe: Daten visualisieren**

Erstellen Sie zwei Darstellungen der Stichprobe:

1. Ein Histogramm zur Darstellung der Verteilung
2. Einen Boxplot zur Visualisierung von Lage und Streuung

Nutzen Sie `matplotlib.pyplot` und stellen Sie die beiden Plots nebeneinander in einem gemeinsamen Plot-Fenster dar.

<!-- grade_id: cell-d6041502ebc7fca4, points: 0, grade: true, solution: true -->
```python

import matplotlib.pyplot as plt

# Plotte Histogramm über data
# Verwende 20 bins
# Die x-Achse soll mit "Wert" gelabelt sein
# Die y-Achse soll das Label "Häufigkeit" tragen
# Der Title des Plots ist "Histogramm der Stichprobe"

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-18fa2f5058cd34d4, points: 0, grade: true, solution: true -->
```python
# Plotte Boxplot über data
# Der Plot soll horizontal (nicht vertikal) sein
# Die x-Achse soll das Label "Wert" tragen
# Der Title des Plots ist "Boxplot der Stichprobe"

# YOUR CODE HERE
raise NotImplementedError()
```