<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Einführung Machine Learning
### Sommersemester 2026
Prof. Dr. Heiner Giefers

# Das Gradientenverfahren bei der linearen Regression

In diesem Notebook wollen wir die lineare Regression mit Hilfe des Gradientenverfahrens implementieren. Dabei werden wir schrittweise das Modell, die Kostenfunktion, die Gradientenberechnung und den Optimierungsalgorithmus als Python-Funktionen formulieren und testen.

## Import und Vorbereitung
Zunächst importieren wir die benötigten Pakete.

```python
import numpy as np
import matplotlib.pyplot as plt
from sklearn.model_selection import train_test_split
from sklearn.preprocessing import MinMaxScaler
```

## Beispiel-Datensatz
Wir erzeugen einen synthetischen Datensatz, der sich gut durch eine lineare Funktion beschreiben lässt. Wir simulieren einen linearen Zusammenhang mit etwas Rauschen:

$$ y = 4 + 3x + \epsilon $$

```python
# Beispielhafte Daten
np.random.seed(42)
m = 100
X = 2 * np.random.rand(m, 1)
y = 4 + 3 * X + np.random.randn(m, 1)

# Bias-Term ergänzen
X_b = np.c_[np.ones((m, 1)), X]  # X0 = 1 für Bias
y = y.reshape(-1, 1)
```

## Modellfunktion
Wir betrachten eine lineare Modellfunktion der Form:
$$ \hat{y} = X\theta $$
**Aufgabe:*** Schreibe eine Funktion `predict(X, theta)`, die für eine gegebene Datenmatrix `X` und Parametervektor `theta` die Vorhersage berechnet.

<!-- grade_id: cell-bc52720cabeff53e, points: 0, grade: false, solution: true -->
```python
def predict(X, theta):
    """Berechnet die Modellvorhersage durch lineare Regression."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-680c232ed6091c6a, points: 0, grade: true, solution: false -->
```python
# Test
theta_test = np.array([[4], [3]])
y_pred = predict(X_b, theta_test)
assert y_pred.shape == (100, 1), 'Falsches Format der Ausgabe'
assert np.allclose(y_pred[0], 4 + 3 * X[0]), 'Fehler in der Vorhersagefunktion'
```

## Kostenfunktion
Als Fehlermaß verwenden wir die mittlere quadratische Abweichung (MSE):
$$ J(\theta) = \frac{1}{2m} \sum_{i=1}^{m} (\hat{y}^{(i)} - y^{(i)})^2 $$
**Aufgabe:** Implementiere eine Funktion `cost(X, y, theta)`, die diese Kosten berechnet.

<!-- grade_id: cell-ba3d998ceae05b11, points: 0, grade: false, solution: true -->
```python
def cost(X, y, theta):
    """Berechnet die mittlere quadratische Abweichung."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-920fcc98a6e4ad0f, points: 0, grade: true, solution: false -->
```python
# Test
assert np.isclose(cost(X_b, y, theta_test), 0.5 * np.mean((y_pred - y)**2))
```

## Gradientenberechnung
Für das Gradientenverfahren benötigen wir den Gradienten der Kostenfunktion:
$$ \nabla J(\theta) = \frac{1}{m} X^T(X\theta - y) $$
**Aufgabe:** Implementiere eine Funktion `compute_gradient(X, y, theta)`, die den Gradienten berechnet.

<!-- grade_id: cell-0545454ce6c125f7, points: 0, grade: false, solution: true -->
```python
def compute_gradient(X, y, theta):
    """Berechnet den Gradienten der Kostenfunktion."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-bb1e6e5c61306780, points: 0, grade: true, solution: false -->
```python
# Test
grad = compute_gradient(X_b, y, theta_test)
assert grad.shape == theta_test.shape
```

## Gradientenverfahren
Jetzt können wir das Gradientenverfahren implementieren. In jeder Iteration berechnen wir den Gradienten und aktualisieren die Parameter:
$$ \theta := \theta - \alpha \cdot \nabla J(\theta) $$
**Aufgabe:** Vervollständige die Funktion `gradient_descent(X, y, alpha, n_iter)`, die die Parameter iterativ verbessert.

<!-- grade_id: cell-809d2eefe3bc1503, points: 0, grade: false, solution: true -->
```python
def gradient_descent(X, y, alpha=0.1, n_iter=100):
    """Führt das Gradientenverfahren durch."""
    m, n = X.shape
    theta = np.zeros((n, 1))
    cost_history = []
    theta_history = []
    for _ in range(n_iter):
        theta_history.append(theta.copy())
        # 1. Berechne den Gradienten
        # 2. Aktualisiere Theta
        # YOUR CODE HERE
        raise NotImplementedError()
        cost_history.append(cost(X, y, theta))
    return theta, cost_history, theta_history
```

<!-- grade_id: cell-6e8145b8d7445d9e, points: 0, grade: true, solution: false -->
```python
# Test
theta_gd, hist, hist_t = gradient_descent(X_b, y, alpha=0.1, n_iter=100)
assert len(hist) == 100
assert theta_gd.shape == (2, 1)
```

## Visualisierung des Kostenverlaufs

```python
plt.plot(hist)
plt.xlabel("Iteration")
plt.ylabel("Kosten")
plt.title("Kostenverlauf beim Gradientenverfahren")
plt.grid()
plt.show()
```

## Modellvisualisierung
Wir zeichnen die Regressionsgerade zusammen mit den Trainingsdaten.

```python
plt.scatter(X[:, 0], y, color="blue", label="Daten")
x_vals = np.linspace(0, 2, 100).reshape(-1, 1)
x_plot = np.c_[np.ones((100, 1)), x_vals]
plt.plot(x_vals, predict(x_plot, theta_gd), color="red", label="Regression")
plt.xlabel("x")
plt.ylabel("y")
plt.legend()
plt.grid()
plt.title("Lineare Regression mit Gradientenabstieg")
plt.show()
```

```python
import numpy as np
import matplotlib.pyplot as plt
from matplotlib.animation import FuncAnimation
from IPython.display import HTML

iterations = 20

# Plot vorbereiten
fig, ax = plt.subplots()
sc = ax.scatter(X, y, label="Daten")
line, = ax.plot([], [], "r-", label="Regression")
ax.set_xlim(0, 2)
ax.set_ylim(y.min() - 1, y.max() + 1)
ax.set_xlabel("x")
ax.set_ylabel("y")
ax.set_title("Lineare Regression – Gradientenschritte")
ax.legend()
ax.grid()

# Animationsfunktion
def update(i):
    t = hist_t[i]
    x_vals = np.linspace(0, 2, 100).reshape(-1, 1)
    x_plot = np.c_[np.ones((100, 1)), x_vals]
    y_vals = predict(x_plot, t)
    line.set_data(x_vals, y_vals)
    ax.set_title(f"Iteration {i+1}")
    return line,

# Animation erzeugen
ani = FuncAnimation(fig, update, frames=iterations, interval=500, blit=True)
plt.close()  # Verhindert Doppelanzeige im Notebook

# Animation anzeigen
HTML(ani.to_jshtml())
```